### Codio Activity 10.5: ACF and PACF Plots for ARMA Models

**Expected Time: 60 Minutes**

**Total Points: 70**

This assignment focuses on using the autocorrelation and partial autocorrelation plots to determine parameters for stationary data.  In general, you will first determine the stationarity of a time series using the Dickey Fuller test (or eyeballing it) and then examine the autocorrelation and partial autocorrelation to identify the parameters for each term.

#### Index

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)
- [Problem 3](#Problem-3)
- [Problem 4](#Problem-4)
- [Problem 5](#Problem-5)
- [Problem 6](#Problem-6)
- [Problem 7](#Problem-7)

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import statsmodels.api as sm
from statsmodels.tsa.arima.model import ARIMA
from statsmodels.tsa.stattools import adfuller
from statsmodels.graphics.tsaplots import plot_acf, plot_pacf
from sklearn.model_selection import train_test_split
import warnings
```

```python
warnings.filterwarnings('ignore')
```

### The Data

Two datasets are used to examine stationarity and autoregression and moving average components for ARMA models.  The first is the recruits data, encountered earlier.  The second is a series of Quarterly GNP data from the United States from 1947 through 2002. In the first you predict the number of recruits and the second your target is the difference of the logarithm of the GNP. 

```python
recruits = pd.read_csv('data/recruitment.csv', index_col=0)
```

```python
recruits.head()
```

```python
plt.plot(recruits.values)
plt.grid()
plt.title('Recruits Data');
```

[Back to top](#-Index)

### Problem 1

#### Is it Stationary? 

**10 Points**

As discussed, our ARMA models are only applicable for stationary data.  Use the `adfuller` function to determine if the recruits data is stationary at the 0.05 level.  Assign your answer as a string to `ans1` below.  

<!-- grade_id: cell-b420b8d4bbddc1da, points: 0, grade: false, solution: true -->
```python
### GRADED

ans1 = ''

### BEGIN SOLUTION
ans1 = 'yes'
### END SOLUTION

# Answer check
print(ans1)
print(type(ans1))
```

<!-- grade_id: cell-7a6c4c6b86a94ada, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans1_ = 'yes'
#
#
#
assert type(ans1_) == type(ans1)
assert ans1_ == ans1
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 2

#### Building train and test set

**10 Points**

Now, we use the familiar `train_test_split` and set `shuffle = False` to create a temporal train and test set.  Leave all arguments to default except `shuffle`.  Assign your results as `y_hist` and `y_future` below. 

<!-- grade_id: cell-361b0197b15eb4b4, points: 0, grade: false, solution: true -->
```python
### GRADED

y_hist, y_future = '', ''

### BEGIN SOLUTION
y_hist, y_future = train_test_split(recruits, shuffle=False)
### END SOLUTION

# Answer check
print("History\n=========")
print(y_hist.tail())
print("Future\n==========")
print(y_future.head())
```

<!-- grade_id: cell-300ec4dd5f8332ed, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
y_hist_, y_future_ = train_test_split(recruits, shuffle=False)
#
#
#
np.testing.assert_array_equal(y_hist, y_hist_)
np.testing.assert_array_equal(y_future, y_future_)
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 3

#### Examining acf and pacf

**10 Points**

Below, the ACF and PACF plots are shown.  While the ACF plot isn't incredibly helpful, the PACF may suggest using a value of `p = 1` in an ARMA model.  As such, create and fit an ARIMA model with `p = 1` and `q = 1`.  Assign your fit model as `arma` below.

```python
fig, ax = plt.subplots(1, 2, figsize = (16, 5))
plot_acf(y_hist, ax = ax[0]);
ax[0].grid()
plot_pacf(y_hist, ax = ax[1], method = 'ywm');
ax[1].grid()
```

```python
y_hist.index = pd.to_datetime(y_hist.index)
```

```python
y_hist.info()
```

<!-- grade_id: cell-563364865387b256, points: 0, grade: false, solution: true -->
```python
### GRADED

arma = ''

### BEGIN SOLUTION
arma = ARIMA(y_hist['value'], order = (1, 0, 1)).fit()
### END SOLUTION

# Answer check
print(arma)
```

<!-- grade_id: cell-a3ea1dd19f75dd51, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
y_hist_, y_future_ = train_test_split(recruits, shuffle=False)
arma_ = ARIMA(y_hist_, order = (1, 0, 1)).fit()
#
#
#
assert type(arma) == type(arma_)
assert list(arma.model_orders.values()) == list(arma_.model_orders.values())
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 4

#### Making Predictions

**10 Points**

Use the `arma` object to make predictions for the training data.  Assign these results as `hist_preds` below.  Uncomment the code to view a plot of the results against the original series. 

<!-- grade_id: cell-206d42551c887e2b, points: 0, grade: false, solution: true -->
```python
### GRADED

hist_preds = ''

### BEGIN SOLUTION
hist_preds = arma.predict()
### END SOLUTION

# Answer check
print(hist_preds.tail())
# plt.figure(figsize = (12, 4))
# plt.plot(hist_preds, label = 'model')
# plt.plot(y_hist, label = 'data')
# plt.legend()
# plt.grid()
# plt.title('Result of ARMA Model');
```

<!-- grade_id: cell-50433e3e804f0829, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
y_hist_, y_future_ = train_test_split(recruits, shuffle=False)
arma_ = ARIMA(y_hist_, order = (1, 0, 1)).fit()
hist_preds_ = arma_.predict()
#
#
#
np.testing.assert_array_equal(hist_preds, hist_preds_)
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 5

#### Forecasting with the ARMA model

**10 Points**

Finally, to use the forecasting capabilities of the model, pass the number of steps to forecast in the future.  Assign the forecast into the future to match up with `y_future` values as `future_preds` below.  

```python
y_future.index = pd.to_datetime(y_future.index)
```

<!-- grade_id: cell-8087b94c5d6e7fb3, points: 0, grade: false, solution: true -->
```python
### GRADED

future_preds = ''

### BEGIN SOLUTION
future_preds = arma.forecast(steps = len(y_future))
### END SOLUTION

# Answer check
print(future_preds.tail())
print(y_future.tail())
```

<!-- grade_id: cell-e963f053833268ae, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
y_hist_, y_future_ = train_test_split(recruits, shuffle=False)
arma_ = ARIMA(y_hist_, order = (1, 0, 1)).fit()
future_preds_ = arma_.forecast(steps = len(y_future))
#
#
#
pd.testing.assert_series_equal(future_preds, future_preds_)
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 6

#### The GNP series

**10 Points**

Below, the `gnp` data is loaded and displayed.  This data is transformed according to the first difference of the logarithm so as to form a stationary series.  Then, the ACF and PACF plots are shown on the stationary series.  These suggest that an AR(2) and MA(2) model might be appropriate.  Build an `ARIMA` model on `y` and predict as `preds`.  Uncomment the code to visualize the predictions.

```python
gnp = pd.read_csv('data/gnp.csv', index_col=0)
gnp.index = pd.Index(pd.date_range("1947-Q1", "2002-Q4", freq = "Q"))
gnp.head()
```

```python
gnp.plot()
```

```python
y = np.log(gnp).diff().dropna()
```

```python
#note the stationarity
adfuller(y)
```

```python
fig, ax = plt.subplots(1, 2, figsize = (14, 4))
plot_acf(y, ax = ax[0]);
ax[0].grid()
plot_pacf(y, ax = ax[1])
ax[1].grid();
```

<!-- grade_id: cell-6422fba5ad9a850a, points: 0, grade: false, solution: true -->
```python
### GRADED

arma2 = ''
preds = ''

### BEGIN SOLUTION
arma2 = ARIMA(y, order = (2, 0, 2)).fit()
preds = arma2.predict()
### END SOLUTION

# # Answer check
# plt.figure(figsize = (12, 4))
# plt.plot(y, label = 'data')
# plt.plot(preds, label = 'predictions')
# plt.legend()
# plt.grid()
# plt.title('Result of ARMA Model');
```

<!-- grade_id: cell-b3ca01ade6e71885, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
arma2_ = ARIMA(y, order = (2, 0, 2)).fit()
preds_ = arma2_.predict()
#
#
#
np.testing.assert_array_equal(preds, preds_)
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 7

#### Errors and Autocorrelation

**10 Points**

Below, subtract the predictions from the actual series.  Determine the stationarity of the results by examining the autocorrelation plot of the residuals.  Is there structure remaining in the series based on this?  Assign your answer as a string to `ans7` below -- 'yes' or 'no'.

```python
preds = pd.DataFrame(preds)
preds.columns = ['value']
```

<!-- grade_id: cell-e42daf6f6975b401, points: 0, grade: false, solution: true -->
```python
### GRADED

resids = ''
ans7 = ''

### BEGIN SOLUTION
resids = y - preds
ans7 = 'no'
### END SOLUTION

# Answer check
plot_acf(resids);
```

<!-- grade_id: cell-984860c93043321d, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
resids_ = y - preds
ans7_ = 'no'
#
#
#
np.testing.assert_array_equal(resids_, resids)
assert type(ans7_) == type(ans7)
assert ans7_ == ans7
### END HIDDEN TESTS
```